# MATH 303 Midterm / Evening session


## Student ID

Enter your student ID number in the cell below using the command:

```
id_student = XXXXXXXX # Replace XXXXXXXX with your student ID number
```

<!-- grade_id: cell-d38091f9b32f40a5, points: 0, grade: false, solution: true -->
```python
id_student = 
```

Run the test in the cell below to verify your student ID number:

<!-- grade_id: cell-11a780af8ba60354, points: 0, grade: true, solution: false -->
```python
assert isinstance(id_student,int) , "Your student ID should be an integer."
print("Your student ID is",id_student)
```

## Problem 1

Run the following cell and note the matrix $P$ and states $i$ and $j$ to be used in Problem 1.

```python
from sympy import latex, Matrix
from IPython.display import display, Math

y = str(id_student); a = 7*id_student + int(y[3]); b = a % 4; i = [1,2,2,1]; j = [3,3,4,4];
P1 = [[[0.0,1.0,0.0,0.0,0.0],[0.4,0.6,0.0,0.0,0.0],[0.5,0.0,0.0,0.5,0.0],[0.0,0.0,0.5,0.0,0.5],[0.0,0.0,0.0,0.0,1.0]],
      [[0.5,0.5,0.0,0.0,0.0],[1.0,0.0,0.0,0.0,0.0],[0.0,0.0,0.0,0.4,0.6],[0.0,0.6,0.4,0.0,0.0],[0.0,0.0,0.0,0.0,1.0]],
      [[0.6,0.4,0.0,0.0,0.0],[1.0,0.0,0.0,0.0,0.0],[0.0,0.5,0.0,0.5,0.0],[0.0,0.0,0.5,0.0,0.5],[0.0,0.0,0.0,0.0,1.0]],
      [[0.0,1.0,0.0,0.0,0.0],[0.5,0.5,0.0,0.0,0.0],[0.0,0.0,0.0,0.4,0.6],[0.0,0.6,0.4,0.0,0.0],[0.0,0.0,0.0,0.0,1.0]]]

display(Math('P = ' + latex(Matrix(P1[b])))); display(Math('i = ' + str(i[b]))); display(Math('j = ' + str(j[b])));
```

## Problem 2

Run the following cell and note the transition probabilities to be used in Problem 2.

```python
from IPython.display import display, Math

y = str(id_student); a = 5*id_student+ int(y[4]); b = a % 4;
P2 = [[1.0,0.8,0.2],[1.0,0.6,0.4],[1.0,0.75,0.25],[1.0,0.9,0.1]]

display(Math('p_{0,0} = ' + str(P2[b][0])))
display(Math('\\text{For all } i>0,\; p_{i,i+1} = ' + str(P2[b][1]) + ', p_{i,i-1} = ' + str(P2[b][2])))
```

## Problem 3

Run the following cell and note the state $i$ to be used in Problem 3.

```python
y = str(id_student); a = id_student + int(y[3]); b = a % 4;

display(Math('i = {}'.format(b + 1)))
```

## Problem 4

Consider the parameters $p$ and $q$ obtained after running the following cell.

```python
from IPython.display import display, Math

y = str(id_student); a = int(y[-1]); b = a % 4;
P4 = [[0.5,0.3],[0.3,0.4],[0.4,0.7],[0.6,0.3]]
p,q = P4[b]

display(Math('p = {}'.format(p))); display(Math('q = {}'.format(q)));
```

Pascal plays a videogame, with two levels to complete successively in order to finish it. At each round, Pascal can either complete the level with probability (w.p.) $p$, or lose w.p. $1-p$. If he loses, he has to start from the beginning and can  either decide to quite the game with probability $q$, or continue playing (w.p. $1-q$).

We consider the Markov chain defined by the following states (and note the order):

* Pascal has quit (state "0")
* Pascal plays the first level (state "1")
* Pascal plays the second level (state "2")
* Pascal has finished the game (state "3")

### Problem 4(a)

Using the values for $p$ and $q$ displayed above, find the transition matrix $P$ associated with the process. Enter the values of the matrix $P$ in the cell below and run the cell.

<!-- grade_id: cell-71593d42d41b3a86, points: 0, grade: false, solution: true -->
```python
P = [[x,x,x,x],[x,x,x,x],[x,x,x,x],[x,x,x,x]] # Complete each entry marked with 'x' of the transition matrix with the appropriate value
```

<!-- grade_id: cell-23122778a5b20901, points: 1, grade: true, solution: false -->
```python
"Verify P has the correct size."
import numpy as np
assert np.array(P).shape == (4,4)
print("Problem 4a Test 1: Success!")
```

<!-- grade_id: cell-f344555fc43f9d2c, points: 1, grade: true, solution: false -->
```python
"Verify rows of P add to 1."
import numpy as np
assert np.allclose(np.array(P).sum(axis=1),[1,1,1,1])
print("Problem 4a Test 2: Success!")
```

<!-- grade_id: cell-49c2de9f3c25fd04, points: 2, grade: true, solution: false -->
```python
"Verify entries of P. This cell contains hidden tests."
import numpy as np
assert np.allclose(np.array(P)[0],[1,0,0,0])
print("Problem 4a Test 3: Success!")
```

### Problem 4(b) 

You saw in week 2 that the mean time spent in a transient state can be found by inverting the following matrix $\mathbf{S}=(\mathbf{I}-\mathbf{P}_T)^{-1}$, where $\mathbf{P}_T$ is the transition matrix of transient states only and $(S)_{ij}$ is the expected number of times where the Markov Chain is in transient state $j$, given that it starts in transient state $i$.

Enter the values of the matrix $\mathbf{P}_T$ in the cell below and run the cell. (Make sure to put the transient states in the correct order.)

<!-- grade_id: cell-4f33345fffb3828f, points: 0, grade: false, solution: true -->
```python
import numpy as np

PT = 

PTmatrix = np.matrix(PT)
S = np.linalg.inv(np.identity(PTmatrix.shape[0]) - PTmatrix)
display(Math('S = ' + latex(Matrix(S))))
```

<!-- grade_id: cell-fe5b03ee82c46adf, points: 1, grade: true, solution: false -->
```python
"Verify PT is a square matrix."
PTmatrix = np.array(PT)
assert PTmatrix.shape[0] == PTmatrix.shape[1]
print("Problem 4b Test 1: Success!")
```

<!-- grade_id: cell-58e9c896dfa8662c, points: 2, grade: true, solution: false -->
```python
"Verify entries of PT. This cell contains hidden tests."
PTmatrix = np.array(PT)
assert PTmatrix.min() >= 0 and PTmatrix.max() <= 1
print("Problem 4b Test 2: Success!")
```

### Problem 4(c)

Use the matrix $\mathbf{S}$ to find the average number of times Pascal plays level 1 until he either quits or finishes the game. Enter this value as `times` in the cell below.

<!-- grade_id: cell-7918991314222c3d, points: 0, grade: false, solution: true -->
```python
times =
```

<!-- grade_id: cell-b386ead2789d4cba, points: 1, grade: true, solution: false -->
```python
"Verify times is the correct value. This cell contains hidden tests."
assert times >= 0
print("Problem 4c Test 1: Success!")
```