# MATH 303 Midterm / morning session

## Student ID

Enter your student ID number in the cell below using the command:

```
id_student = XXXXXXXX # Replace XXXXXXXX with your student ID number
```

<!-- grade_id: cell-48907aac1685dbe7, points: 0, grade: false, solution: true -->
```python
id_student = 
```

Run the test in the cell below to verify your student ID number:

<!-- grade_id: cell-720ebd2f7ed7634a, points: 0, grade: true, solution: false -->
```python
assert isinstance(id_student,int) , "Your student ID should be an integer."
print("Your student ID is",id_student)
```

## Problem 1

Run the following cell and note the matrix $P$ and state $i$ to be used in Problem 1.

```python
from sympy import latex, Matrix
from IPython.display import display, Math

y = str(id_student); a = id_student + int(y[2]); b = a % 4; i = [5,3,4,2];
P1 = [[[0.0,0.0,0.0,1.0,0.0],[0.0,0.5,0.0,0.0,0.5],[0.2,0.0,0.6,0.0,0.2],[1.0,0.0,0.0,0.0,0.0],[0.0,0.4,0.0,0.0,0.6]],
      [[0.8,0.0,0.2,0.0,0.0],[0.0,0.0,0.0,1.0,0.0],[0.4,0.0,0.6,0.0,0.0],[0.0,1.0,0.0,0.0,0.0],[0.0,0.0,0.1,0.1,0.8]],
      [[0.2,0.6,0.2,0.0,0.0],[0.0,0.0,0.0,0.0,1.0],[0.0,0.0,0.2,0.8,0.0],[0.0,0.0,0.4,0.6,0.0],[0.0,1.0,0.0,0.0,0.0]],
      [[0.4,0.6,0.0,0.0,0.0],[0.5,0.5,0.0,0.0,0.0],[0.0,0.0,0.0,0.0,1.0],[0.3,0.0,0.0,0.3,0.4],[0.0,0.0,1.0,0.0,0.0]]]

display(Math('P = ' + latex(Matrix(P1[b])))); display(Math('i = ' + str(i[b])));
```

## Problem 2

Run the following cell and note the transition probabilities to be used in Problem 2.

```python
from IPython.display import display, Math

y = str(id_student); a = id_student + int(y[4]); b = a % 4;
P2 = [[0.25,0.75,0.25,0.5,0.25],[0.5,0.5,0.5,0.25,0.25],[0.2,0.8,0.2,0.4,0.4],[0.6,0.4,0.6,0.2,0.2]]

display(Math('p_{0,1} = ' + str(P2[b][0]) + ', p_{0,0} = ' + str(P2[b][1])))
display(Math('\\text{For all } i>0,\; p_{i,i+1} = ' + str(P2[b][2]) + ', p_{i,i-1} = ' + str(P2[b][3]) + ' \\text{ and } p_{i,i} = ' + str(P2[b][4])))
```

## Problem 3

Run the following cell and note the state $i$ to be used in Problem 3.

```python
y = str(id_student); a = id_student + int(y[3]); b = a % 4

display(Math('i = {}'.format(b + 1)))
```

## Problem 4

Consider the parameters $p,q$ and $r$ obtained after running the following cell.

```python
from IPython.display import display, Math

y = str(id_student); a = int(y[-1]); b = a % 4;
P4 = [[0.5,0.3,0.2],[0.3,0.5,0.7],[0.4,0.7,0.5],[0.6,0.5,0.4]]
p,q,r = P4[b]

display(Math('p = {}'.format(p))); display(Math('q = {}'.format(q))); display(Math('r = {}'.format(r)));
```

A student makes trips between UBC and their home. If it rains (r), the student takes the bus (B). If it is sunny (s), the student can either bike (b) with probability $p$ or take the bus (with probability $1-p$). If the weather is sunny, the probability that it is still sunny on the next trip is $q$ and if it is raining, the probability that it still rains on the next trip is $r$.

We consider the Markov chain defined by the joint observation of the weather and the transport used by the student: Since the student only takes the bus when it rains, there are three states: $Br$ (it rains and the student takes a bus), $Bs$ (it is sunny and the student takes the bus), and $bs$ (it is sunny and the student bikes)

### Problem 4(a)

Using the values for $p,q$ and $r$ displayed above, find the transition matrix $P$ associated with the process. Enter the values of the matrix $P$ in the cell below and run it. Note the order of the states:

* State 1: Rain and student takes the bus (Br)
* State 2: Sunny and student takes the bus (Bs)
* State 3: Sunny and student rides a bike (bs)

Enter the values of $P$ in the cell below:

<!-- grade_id: cell-44006d507bd32821, points: 0, grade: false, solution: true -->
```python
P = [[x,x,x],[x,x,x],[x,x,x]] # Complete each entry marked with 'x' of the transition matrix with the appropriate value
```

<!-- grade_id: cell-55dc83f779966789, points: 1, grade: true, solution: false -->
```python
"Verify P has the correct size."
import numpy as np
assert np.array(P).shape == (3,3)
print("Problem 4a Test 1: Success!")
```

<!-- grade_id: cell-b1fa435fe6b8dc2f, points: 1, grade: true, solution: false -->
```python
"Verify rows of P add to 1."
import numpy as np
assert np.allclose(np.array(P).sum(axis=1),[1,1,1])
print("Problem 4a Test 2: Success!")
```

<!-- grade_id: cell-4849f858c48d605e, points: 2, grade: true, solution: false -->
```python
"Verify entries of P. This cell contains hidden tests."
import numpy as np
assert np.allclose(np.array(P).sum(axis=1),[1,1,1])
print("Problem 4a Test 3: Success!")
```

### Question 4(b)

Run the following cell. This should give you an empirical estimate of the stationary distribution $\mu$ (we assume it exists).

```python
import numpy as np

class MarkovChain(object):
    def __init__(self, transition_matrix, states):
        self.transition_matrix = np.atleast_2d(transition_matrix)
        self.states = states
        self.index_dict = {self.states[index]: index for index in range(len(self.states))}
        self.state_dict = {index: self.states[index] for index in range(len(self.states))}
 
    def next_state(self, current_state):
        return np.random.choice(self.states,p=self.transition_matrix[self.index_dict[current_state], :])
 
    def generate_states(self, current_state, no=10):
        future_states = []
        for i in range(no):
            next_state = self.next_state(current_state)
            future_states.append(next_state)
            current_state = next_state
        return future_states
    
def get_empirical_dist(states, transition_matrix, initial_state, nrealizations, N):
    chain = MarkovChain(transition_matrix=transition_matrix, states=states)
    arr = np.zeros(((N+1), nrealizations))
    for r in range(nrealizations):
        simulation = chain.generate_states(current_state=initial_state, no=(N+1))
        arr[:, r] = np.array(simulation)
        nstates = len(states)
        count = np.zeros(nstates)
    for i in range(nstates):
        state = int(states[i])
        index = np.where(arr[N] == state)
        count[i] = len(index[0])        
        distribution = count/nrealizations
    return distribution

# Let -1 denote 'br', 0 denote 'Bs' and 1 denote 'bs'.
empirical_distribution = get_empirical_dist(states=["-1", "0", "1"], transition_matrix=P, initial_state="-1", nrealizations=1000, N=10)

print('The empirical stationary distribution is',empirical_distribution)
```

Assuming that the distribution of the first trip followed $\mu$ and using results from Problem 4(a), enter below (as `proba`) the probability that given that a student takes a bus on a trip, they biked during the previous trip.

<!-- grade_id: cell-2b6c9c8bdea10cc3, points: 0, grade: false, solution: true -->
```python
proba = 
```

<!-- grade_id: cell-64d86635b1bbd3ce, points: 1, grade: true, solution: false -->
```python
"Verify proba is a number between 0 and 1."
assert proba <= 1 and proba >= 0 , "Probability should be between 0 and 1."
print("Problem 4b Test 1: Success!")
```

<!-- grade_id: cell-a01e1dc729be8742, points: 0, grade: true, solution: false -->
```python
"Verify proba has the correct value. This cell contains hidden tests."
assert proba <= 1
print("Problem 4b Test 2: Success!")
```